<a href="https://colab.research.google.com/github/mongarejoshua/ENG-219-084-2023/blob/main/Copy_of_Assignment1_Control_Statements.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 02.1 (if-else)

Consider the following assessment criteria which map a score out of 100 to an
assessment grade:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| Excellent        | $\ge 85$               |
| Very good        | $\ge 76.5$ and $< 85$  |
| Good             | $\ge 64$ and $< 76.5$  |
| Need improvement | $\ge 40$ and $< 64$    |
| Did you try?     | $< 40$                 |

Write a program that, given an a score, prints the appropriate grade. Print an error message if the input score is greater than 100 or less than zero.

<!-- grade_id: cell-71121f470be22e4a, points: 0, grade: false, solution: true -->
```python
# Score from user
score = 209

#my code
if score > 100 or score < 0:
    print("Error, score must be between 0 and 100.")
elif score >= 85:
    print('Excellent')
elif score >= 76.5:
    print('Very good')
elif score >= 64:
    print('Good')
elif score >= 40:
    print('Need improvement')
else:
    print('Did you try?')
```

Output:
```
Error, score must be between 0 and 100.
```

## Exercise 02.2 (bisection)

Bisection is an iterative method for finding approximate roots of a function. Say we know that the function $f(x)$ has one root between $x_{0}$ and $x_{1}$ ($x_{0} < x_{1}$). We then:

- Evaluate $f$ at the midpoint $x_{\rm mid} = (x_0 + x_1)/2$, i.e. compute
   $f_{\rm mid} = f(x_{\rm mid})$
- Evaluate $f(x_0) \cdot f(x_{\rm mid})$

  - if $f(x_0) \cdot f(x_{\rm mid}) < 0$:

    $f$ must change sign somewhere between $x_0$ and $x_{\rm mid}$, hence the root must lie between
    $x_0$ and $x_{\rm mid}$, so set $x_1 = x_{\rm mid}$.
   
  - else:

    $f$ must change sign somewhere between $x_{\rm mid}$ and $x_1$, so set
    $x_0 = x_{\rm mid}$.

The above steps can be repeated a specified number of times, or until $|f_{\rm mid}|$
is below a tolerance, with $x_{\rm mid}$ being the approximate root.


### Task

The function

$$
f(x) = \frac{x^{5}}{10} + x^3 - 10x^2 + 4x + 7
$$


has one root in the range $0 < x < 2$.

1. Use the bisection method to find an approximate root $x_{r}$ using 20 iterations
   (use a `for` loop).
2. Use the bisection method to find an approximate root $x_{r}$ such that
   $\left| f(x_{r}) \right| < 1 \times 10^{-6}$ and report the number of iterations
   required (use a `while` loop).

Store the approximate root using the variable `x_mid`, and store $f(x_{\rm mid})$ using the variable `f`.

*Hint:* Use  `abs` to compute the absolute value of a number, e.g. `y = abs(x)` assigns the absolute value of `x` to `y`.

#### (1) Using a `for` loop.

<!-- grade_id: cell-3a0d0a59613d6b3f, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

for n in range(20):
    # midpoint
    x_mid = (x0 + x1) / 2

    # Evaluate function at (i) left end-point and at (ii) midpoint
    f0 = (x0**5) / 10 + x0**3 - 10 * x0**2 + 4 * x0 + 7
    f = (x_mid**5) / 10 + x_mid**3 - 10 * x_mid**2 + 4 * x_mid + 7

    if f0 * f < 0:
        x1 = x_mid
    else:
        x0 = x_mid


    print(f"Interation {n}: x_mid: {x_mid}, f: {f}")
```

Output:
```
Interation 0: x_mid: 1.0, f: 2.0999999999999996
Interation 1: x_mid: 1.5, f: -5.365625000000001
Interation 2: x_mid: 1.25, f: -1.36669921875
Interation 3: x_mid: 1.125, f: 0.4477813720703132
Interation 4: x_mid: 1.1875, f: -0.4408627510070797
Interation 5: x_mid: 1.15625, f: 0.008327355980872753
Interation 6: x_mid: 1.171875, f: -0.21507753478363156
Interation 7: x_mid: 1.1640625, f: -0.1030741602036862
Interation 8: x_mid: 1.16015625, f: -0.04729774910292761
Interation 9: x_mid: 1.158203125, f: -0.019466230897839054
Interation 10: x_mid: 1.1572265625, f: -0.005564689502373099
Interation 11: x_mid: 1.15673828125, f: 0.0013825210451265946
Interation 12: x_mid: 1.156982421875, f: -0.0020907873792292975
Interation 13: x_mid: 1.1568603515625, f: -0.00035405894194262544
Interation 14: x_mid: 1.15679931640625, f: 0.0005142496094645566
Interation 15: x_mid: 1.156829833984375, f: 8.009997302949046e-05
Interation 16: x_mid: 1.1568450927734375, f: -0.00013697832466341708
Interation 17: x_mid: 1.
```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
assert math.isclose(x_mid, 1.1568355560302734)
assert abs(f) < 1e-5
```

#### (2) Using a `while` loop

Use the variable `counter` for the iteration number.

*Remember to guard against infinite loops.*

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

tol = 1.0e-6
error = tol + 1.0

# Iterate until tolerance is met
counter = 0
while error > tol:
    x_mid = (x0 + x1) / 2

    f = (x_mid**5) / 10 + x_mid**3 - 10 * x_mid**2 + 4 * x_mid + 7
    f0 = (x0**5) / 10 + x0**3 - 10 * x0**2 + 4 * x0 + 7

    error = abs(f)

    if f0 * f < 0:
        x1 = x_mid
    else:
        x0 = x_mid

    counter += 1

    # Guard against an infinite loop
    if counter > 1000:
        print("Oops, iteration count is very large. Breaking out of while loop.")
        break

    print(counter, x_mid, error)
    print(f"Counter: {counter}, x_mid: {x_mid}, Error: {error}")
```

Output:
```
1 1.0 2.0999999999999996
Counter: 1, x_mid: 1.0, Error: 2.0999999999999996
2 1.5 5.365625000000001
Counter: 2, x_mid: 1.5, Error: 5.365625000000001
3 1.25 1.36669921875
Counter: 3, x_mid: 1.25, Error: 1.36669921875
4 1.125 0.4477813720703132
Counter: 4, x_mid: 1.125, Error: 0.4477813720703132
5 1.1875 0.4408627510070797
Counter: 5, x_mid: 1.1875, Error: 0.4408627510070797
6 1.15625 0.008327355980872753
Counter: 6, x_mid: 1.15625, Error: 0.008327355980872753
7 1.171875 0.21507753478363156
Counter: 7, x_mid: 1.171875, Error: 0.21507753478363156
8 1.1640625 0.1030741602036862
Counter: 8, x_mid: 1.1640625, Error: 0.1030741602036862
9 1.16015625 0.04729774910292761
Counter: 9, x_mid: 1.16015625, Error: 0.04729774910292761
10 1.158203125 0.019466230897839054
Counter: 10, x_mid: 1.158203125, Error: 0.019466230897839054
11 1.1572265625 0.005564689502373099
Counter: 11, x_mid: 1.1572265625, Error: 0.005564689502373099
12 1.15673828125 0.0013825210451265946
Counter: 12, x_mid: 1.15673828125, Err
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
## tests ##
assert counter == 24
assert abs(f) < 1.0e-6
```

## Exercise 02.3 (series expansion)

For $|x| < 1$ the series:

$$
(1 + x)^{-1/2} = \sum_{n = 0}^{\infty} \frac{(-1)^n (2n)!}{4^n (n!)^2} x^n
$$

converges.

1. Using a `for` statement, approximate $1/\sqrt{0.16}$ using 30 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $1/\sqrt{0.16}$ to within $1 \times 10^{-5}$.

Store the absolute value of the error in the variable `error`.

### Hints

To compute the factorial, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
```
You only need `import math` once at the top of your program. Standard modules, like `math`, will be explained in a later

<!-- The power series expansion for the sine function is:

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$.

Store the absolute value of the error in the variable `error`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the error. -->

#### (1) Using a `for` loop

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
import math

x = -0.84

approx_f = 0.0 #approx. function

true = 1/math.sqrt(0.16)

for n in range(31):
    numerator = (-1)**n * math.factorial(2*n) * x**n
    denominator = 4**n * (math.factorial(n))**2
    approx_f += numerator / denominator

error = abs(true - approx_f)

print(f"The error is:\n{error}")
```

Output:
```
The error is:
0.0026433650090669225
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
## test ##
assert error < 1.0e-2
```

#### (2) Using a `while` loop

*Remember to guard against infinite loops.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
import math

x = -0.84

tol = 1.0e-5
error = tol + 1.0

approx_f = 0.0
n = 0

while error > tol and n < 1000:
    factorial = 1
    for i in range(1, 2*n + 2):
        factorial *= i

    numerator = (-1)**n * math.factorial(2*n) * x**n
    denominator = 4**n * (math.factorial(n))**2

    approx_f += numerator / denominator

    error = abs(true - approx_f)

    # Increment counter
    n += 1

print("\nThe error is:", error)
print("Number of terms in series:", n)
```

Output:
```
The error is: 8.689878570500298e-06
Number of terms in series: 62
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
## test ##
assert error <= 1.0e-5
```